--- 
title_meta  : บทที่ 1 
title       : แนะนำการใช้งานเบื้องต้น
description : "ในบทเรียนนี้คุณณจะได้เริ่มต้นใช้งานภาษา R คุณจะได้เรียนรู้การใช้งาน console เป็นเครื่องคิดเลข และการใส่ค่าให้กับตัวแปร นอกจากนี้คุณจะได้รู้จักกับชนิดของข้อมูลพื้นฐานในภาษา R เรามาเริ่มเรียนกันเถอะ!"

--- type:NormalExercise xp:100 skills:1 key:15d729634a
## เริ่มต้นใช้งาน

คุณจะต้องพิมพ์โค้ดภาษา R ลงใน editor ทางด้านขวาเพื่อตอบแบบฝึกหัด เมื่อคุณกดปุ่ม 'Submit Answer' โค้ดทุกๆบรรทัดจะถูกประมวลผล และคุณจะได้รับข้อความว่าโค้ดของคุณนั้นถูกต้องหรือไม่ ผลลัพธ์ของโค้ดจะถูกแสดงอยู่ใน console ที่มุมขวาล่าง

R ใช้เครื่องหมาย `#` ในการคอมเมนต์เพื่ออธิบายการทำงานของโค้ด คอมเมนต์จะไม่ถูกประมวลผลเหมือนกับโค้ด และส่งผลใดๆต่อผลลัพธ์ ตัวอย่างเช่น _Calculate 3 + 4_ ใน editor ด้านขวานั้น ถือว่าเป็นคอมเมนต์

คุณสามารถประมวลผลคำสั่ง R ผ่านทางทาง Console ได้โดยตรงอีกด้วย ซึ่งเป็นช่องทางที่ดีสำหรับการทดลองคำสั่งต่างๆ เนื่องจากคำสั่งที่ป้อนเข้าไปจะไม่ถูกนำมาตรวจสอบความถูกต้อง

*** =instructions
- ใน editor ทางด้านขวา จะมีตัวอย่างโค้ดมาให้ คุณสามารถบอกได้หรือไม่ว่าบรรทัดไหนเป็นโค้ดภาษา R จริงๆ และบรรทัดไหนเป็นคอมเมนต์?
- เขียนโค้ดเพิ่มอีกหนึ่งบรรทัดเพื่อคำนวณค่าผลรวมของ 6 และ 12 จากนั้นกดปุ่ม 'Submit Answer'

*** =hint
แค่เขียนโค้ดภาษา R อีกหนึ่งบรรทัด เพื่อคำนวณผลรวมของ 6 และ 12 เหมือนตัวอย่างที่ให้ไว้ในโค้ด!

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Calculate 3 + 4
3 + 4

# Calculate 6 + 12

```

*** =solution
```{r}
# Calculate 3 + 4
3 + 4

# Calculate 6 + 12
6 + 12
```

*** =sct
```{r}

incorrect_msg <- paste("&#3605;&#3619;&#3623;&#3592;&#3626;&#3629;&#3610;&#3651;&#3627;&#3657;&#3649;&#3609;&#3656;&#3651;&#3592;&#3623;&#3656;&#3634;&#3588;&#3640;&#3603;&#3652;&#3604;&#3657;&#3614;&#3636;&#3617;&#3614;&#3660;", "`6 + 12`", "&#3621;&#3591;&#3651;&#3609;&#3610;&#3619;&#3619;&#3607;&#3633;&#3604;&#3651;&#3627;&#3617;&#3656;&#32;&#3629;&#3618;&#3656;&#3634;&#3648;&#3619;&#3636;&#3656;&#3617;&#3605;&#3657;&#3609;&#3610;&#3619;&#3619;&#3607;&#3633;&#3604;&#3604;&#3657;&#3623;&#3618;&#3648;&#3588;&#3619;&#3639;&#3656;&#3629;&#3591;&#3627;&#3617;&#3634;&#3618;", "`#`", "&#3652;&#3617;&#3656;&#3648;&#3594;&#3656;&#3609;&#3609;&#3633;&#3657;&#3609;&#3650;&#3588;&#3657;&#3604;&#32;&#82;&#32;&#3586;&#3629;&#3591;&#3588;&#3640;&#3603;&#3592;&#3632;&#3652;&#3617;&#3656;&#3606;&#3641;&#3585;&#3611;&#3619;&#3632;&#3617;&#3623;&#3621;&#3612;&#3621;&#33;")
# ตรวจสอบให้แน่ใจว่าคุณได้พิมพ์ `6 + 12` ลงในบรรทัดใหม่ อย่าเริ่มต้นบรรทัดด้วยเครื่องหมาย `#` ไม่เช่นนั้นโค้ด R ของคุณจะไม่ถูกประมวลผล!

success_msg <- " &#3648;&#3618;&#3637;&#3656;&#3618;&#3617;&#3652;&#3611;&#3648;&#3621;&#3618;&#33;&#32;&#3626;&#3633;&#3591;&#3648;&#3585;&#3605;&#3627;&#3619;&#3639;&#3629;&#3652;&#3617;&#3656;&#3623;&#3656;&#3634;&#32;&#99;&#111;&#110;&#115;&#111;&#108;&#101;&#32;&#3649;&#3626;&#3604;&#3591;&#3612;&#3621;&#3586;&#3629;&#3591;&#3650;&#3588;&#3657;&#3604;&#3607;&#3637;&#3656;&#3588;&#3640;&#3603;&#3626;&#3656;&#3591;&#3617;&#3634;&#3629;&#3618;&#3656;&#3634;&#3591;&#3652;&#3619;&#63;&#32;&#3648;&#3617;&#3639;&#3656;&#3629;&#3588;&#3640;&#3603;&#3588;&#3640;&#3657;&#3609;&#3648;&#3588;&#3618;&#3585;&#3633;&#3610;&#3627;&#3609;&#3657;&#3634;&#3605;&#3634;&#3649;&#3621;&#3632;&#3585;&#3634;&#3619;&#3651;&#3594;&#3657;&#3591;&#3634;&#3609;&#3610;&#3607;&#3648;&#3619;&#3637;&#3618;&#3609;&#3649;&#3621;&#3657;&#3623;&#32;&#3648;&#3619;&#3634;&#3617;&#3634;&#3648;&#3619;&#3636;&#3656;&#3617;&#3651;&#3594;&#3657;&#3591;&#3634;&#3609;&#32;&#82;&#32;&#3585;&#3633;&#3609;&#3592;&#3619;&#3636;&#3591;&#3654;&#3648;&#3621;&#3618;&#3604;&#3637;&#3585;&#3623;&#3656;&#3634;&#33;&#32;"
# เยี่ยมไปเลย! สังเกตหรือไม่ว่า console แสดงผลของโค้ดที่คุณส่งมาอย่างไร? เมื่อคุณคุ้นเคยกับหน้าตาและการใช้งานบทเรียนแล้ว เรามาเริ่มใช้งาน R กันจริงๆเลยดีกว่า!

test_output_contains("18", incorrect_msg = incorrect_msg)
success_msg(success_msg)
```

--- type:NormalExercise xp:100 skills:1 key:720745eda5
## การคำนวณทางคณิตศาสตร์ด้วย R 


R สามารถนำมาใช้เป็นเครื่องคิดเลขอย่างง่ายๆได้ พิจารณาตัวดำเนินการทางคณิตศาสตร์ (arithmetic operators) ดังต่อไปนี้:

- บวก (addition): `+`
- ลบ (subtraction): `-`
- คูณ (multiplication): `*`
- หาร (division): `/`
- ยกกำลัง (exponentiation): `^`
- มอดุโล (modulo): `%%`

ตัวดำเนินการสองตัวสุดท้ายอาจจะต้องขยายความสักหน่อย:

- ตัวดำเนินการ `^` จะนำเลขฐานที่อยู่ทางด้านซ้ายของมัน มายกกำลังด้วยเลขชี้กำลังที่อยู่ทางด้านขวาของมัน เช่น `3^2` ได้ 9
- ตัวดำเนินการมอดุโล `%%` เป็นการหาเศษเหลือจากการหารตัวเลขทางด้านซ้ายด้วยตัวเลขทางด้านขวาของมัน เช่น 5 มอดุโล 3 หรือ `5 %% 3` ได้ 2

ด้วยความรู้ข้างต้น จงทำแบบฝึกหัดตามคำสั่งด้านล่างนี้

*** =instructions
- พิมพ์ `2^5` ใน editor เพื่อคำนวณ 2 ยกกำลัง 5
- พิมพ์ `28 %% 6` เพื่อคำนวณ 28 มอดุโล 6
- กด 'Submit Answer' และดูผลลัพธ์ใน Console
- สังเกตว่าสัญลักษณ์ `#` ถูกนำมาใช้ในการเขียนคอมเมนต์ในโค้ดอย่างไร 

*** =hint
อีกตัวอย่างของตัวดำเนินการมอดุโล: `9 %% 2` เท่ากับ `1`

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# บวก
5 + 5 

# ลบ
5 - 5 

# คูณ
3 * 5

# หาร
(5 + 5) / 2 

# ยกกำลัง


# มอดุโล

```

*** =solution
```{r}
# addition
5 + 5 

# subtraction
5 - 5 

# multiplication
3 * 5

# division
(5 + 5) / 2 

# exponentiation
2 ^ 5

# modulo
28 %% 6
```

*** =sct
```{r}
delete_msg <- "&#3629;&#3618;&#3656;&#3634;&#3621;&#3610;&#3605;&#3633;&#3623;&#3629;&#3618;&#3656;&#3634;&#3591;&#3607;&#3637;&#3656;&#3651;&#3627;&#3657;&#3617;&#3634;&#3629;&#3629;&#3585;&#3652;&#3611;&#33;"
# อย่าลบตัวอย่างที่ให้มาออกไป!

exponent_msg <- paste("&#3605;&#3633;&#3623;&#3629;&#3618;&#3656;&#3634;&#3591;&#3585;&#3634;&#3619;&#3618;&#3585;&#3585;&#3635;&#3621;&#3633;&#3591;&#3618;&#3633;&#3591;&#3652;&#3617;&#3656;&#3606;&#3641;&#3585;&#3605;&#3657;&#3629;&#3591;&#32;&#3614;&#3636;&#3617;&#3614;&#3660;", "`2 ^ 5`", "&#3621;&#3591;&#3651;&#3609;&#3610;&#3619;&#3619;&#3607;&#3633;&#3604;&#3651;&#3627;&#3617;&#3656;")
# ตัวอย่างการยกกำลังยังไม่ถูกต้อง พิมพ์ `2 ^ 5` ลงในบรรทัดใหม่

modulo_msg <- paste("&#3604;&#3641;&#3648;&#3627;&#3617;&#3639;&#3629;&#3609;&#3592;&#3632;&#3617;&#3637;&#3611;&#3633;&#3597;&#3627;&#3634;&#3585;&#3633;&#3610;&#3605;&#3633;&#3623;&#3629;&#3618;&#3656;&#3634;&#3591;&#3585;&#3634;&#3619;&#3651;&#3594;&#3657;&#3617;&#3629;&#3604;&#3640;&#3650;&#3621;&#32;&#3614;&#3636;&#3617;&#3614;&#3660;", "`28 %% 6`", "&#3621;&#3591;&#3651;&#3609;&#3610;&#3619;&#3619;&#3607;&#3633;&#3604;&#3651;&#3627;&#3617;&#3656;")
# ดูเหมือนจะมีปัญหากับตัวอย่างการใช้มอดุโล พิมพ์ `28 %% 6` ลงในบรรทัดใหม่

success_msg <- "&#3648;&#3618;&#3637;&#3656;&#3618;&#3617;&#33;&#32;&#3652;&#3611;&#3649;&#3610;&#3610;&#3613;&#3638;&#3585;&#3627;&#3633;&#3604;&#3605;&#3656;&#3629;&#3652;&#3611;&#3652;&#3604;&#3657;&#3648;&#3621;&#3618;&#33;"
# เยี่ยม! ไปแบบฝึกหัดต่อไปได้เลย!

lapply(c("5+5", "5-5", "3*5", "(5+5)/2"), test_output_contains, incorrect_msg = delete_msg)
test_output_contains("2^5", incorrect_msg = exponent_msg)
test_output_contains("28 %% 6", incorrect_msg = modulo_msg)
success_msg(success_msg)
```

--- type:NormalExercise xp:100 skills:1 key:5f200ffd43
## การกำหนดค่าตัวแปร

ในการเขียนโปรแกรม เรามี concept พื้นฐานที่เรียกว่า **ตัวแปร (variable)**

คุณสามารถเก็บค่าต่างๆ (เช่น 4) หรือวัตถุ (เช่น ฟังก์ชันการคำนวณต่างๆ) ไว้ในตัวแปรได้ ซึ่งค่าเหล่านี้จะถูกเรียกมาใช้ภายหลังได้โดยการเรียกชื่อของตัวแปรนั้นๆที่คุณได้กำหนดค่าเอาไว้

คุณสามารถกำหนดค่าของตัวแปร `my_var` ให้เป็น 4 ได้ด้วยคำสั่ง 

```
my_var <- 4
```

*** =instructions
ลองกำหนดค่าของตัวแปร `x` ใน editor ให้เป็น 42 จากนั้นคลิก 'Submit Answer' คุณจะเห็นได้ว่ามีเลข 42 ปรากฎขึ้นมาใน Console เมื่อคุณสั่งให้ R แสดงค่า `x`

*** =hint
คุณสามารถใช้คำสั่ง `my_var <- 4` เป็นตัวอย่างได้ ซึ่งเป็นการกำหนดค่าของตัวแปร `my_var` ให้เป็น 4 คราวนี้ลองกำหนดค่าตัวแปร `x` ให้เป็น 42 ดูบ้าง

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Assign the value 42 to x 
x <- 

# Print out the value of the variable x
x
```

*** =solution
```{r}
# Assign the value 42 to x
x <- 42

# Print out the value of the variable x
x
```

*** =sct
```{r}
undefined_msg <- paste("&#3588;&#3640;&#3603;&#3618;&#3633;&#3591;&#3652;&#3617;&#3656;&#3652;&#3604;&#3657;&#3585;&#3635;&#3627;&#3609;&#3604;&#3588;&#3656;&#3634;&#3651;&#3627;&#3657;&#3585;&#3633;&#3610;&#3605;&#3633;&#3623;&#3649;&#3611;&#3619;", "`x`") 
# คุณยังไม่ได้กำหนดค่าให้กับตัวแปร `x`

incorrect_msg <- paste("&#3588;&#3656;&#3634;&#3607;&#3637;&#3656;&#3585;&#3635;&#3627;&#3609;&#3604;&#3651;&#3627;&#3657;&#3585;&#3633;&#3610;&#3605;&#3633;&#3623;&#3649;&#3611;&#3619;", "`x`", "&#3652;&#3617;&#3656;&#3606;&#3641;&#3585;&#3605;&#3657;&#3629;&#3591;")
# ค่าที่กำหนดให้กับตัวแปร `x` ไม่ถูกต้อง

success_msg <- paste("&#3607;&#3635;&#3652;&#3604;&#3657;&#3604;&#3637;&#3617;&#3634;&#3585;&#33;&#32;&#3626;&#3633;&#3591;&#3648;&#3585;&#3605;&#3652;&#3627;&#3617;&#3623;&#3656;&#3634;&#32;&#82;&#32;&#3652;&#3617;&#3656;&#3652;&#3604;&#3657;&#3649;&#3626;&#3604;&#3591;&#3588;&#3656;&#3634;&#3586;&#3629;&#3591;&#3605;&#3633;&#3623;&#3649;&#3611;&#3619;&#3629;&#3629;&#3585;&#3617;&#3634;&#3605;&#3629;&#3609;&#3607;&#3637;&#3656;&#3588;&#3640;&#3603;&#3585;&#3635;&#3627;&#3609;&#3604;&#3588;&#3656;&#3634;&#3651;&#3627;&#3657;&#3585;&#3633;&#3610;&#3605;&#3633;&#3623;&#3649;&#3611;&#3619;&#3609;&#3633;&#3657;&#3609;&#3654;&#63;", "`x <- 42`", "&#3652;&#3617;&#3656;&#3652;&#3604;&#3657;&#3607;&#3635;&#3651;&#3627;&#3657;&#32;&#82;&#32;&#3649;&#3626;&#3604;&#3591;&#3588;&#3656;&#3634;&#3651;&#3604;&#3654;&#3629;&#3629;&#3585;&#3617;&#3634;&#32;&#3648;&#3614;&#3619;&#3634;&#3632;&#32;&#82;&#32;&#3588;&#3636;&#3604;&#3623;&#3656;&#3634;&#3588;&#3640;&#3603;&#3592;&#3632;&#3609;&#3635;&#3588;&#3656;&#3634;&#3586;&#3629;&#3591;&#3605;&#3633;&#3623;&#3649;&#3611;&#3619;&#3609;&#3637;&#3657;&#3652;&#3611;&#3651;&#3594;&#3657;&#3651;&#3609;&#3629;&#3609;&#3634;&#3588;&#3605;&#32;&#3604;&#3633;&#3591;&#3609;&#3633;&#3657;&#3609;&#3592;&#3638;&#3591;&#3618;&#3633;&#3591;&#3652;&#3617;&#3656;&#3617;&#3637;&#3588;&#3623;&#3634;&#3617;&#3592;&#3635;&#3648;&#3611;&#3655;&#3609;&#3651;&#3604;&#3654;&#3607;&#3637;&#3656;&#3605;&#3657;&#3629;&#3591;&#3649;&#3626;&#3604;&#3591;&#3588;&#3656;&#3634;&#3586;&#3629;&#3591;&#3605;&#3633;&#3623;&#3649;&#3611;&#3619;&#3629;&#3629;&#3585;&#3617;&#3634;&#32;&#3648;&#3619;&#3634;&#3652;&#3611;&#3604;&#3641;&#3649;&#3610;&#3610;&#3613;&#3638;&#3585;&#3627;&#3633;&#3604;&#3605;&#3656;&#3629;&#3652;&#3611;&#3585;&#3633;&#3609;&#3648;&#3621;&#3618;&#33;")
# ทำได้ดีมาก! สังเกตไหมว่า R ไม่ได้แสดงค่าของตัวแปรออกมาตอนที่คุณกำหนดค่าให้กับตัวแปรนั้นๆ? `x <- 42` ไม่ได้ทำให้ R แสดงค่าใดๆออกมา เพราะ R คิดว่าคุณจะนำค่าของตัวแปรนี้ไปใช้ในอนาคต ดังนั้นจึงยังไม่มีความจำเป็นใดๆที่ต้องแสดงค่าของตัวแปรออกมา เราไปดูแบบฝึกหัดต่อไปกันเลย!

test_object("x", undefined_msg = undefined_msg, incorrect_msg = incorrect_msg) 
success_msg(success_msg)
```


--- type:NormalExercise xp:100 skills:1 key:c5944b90eb
## การกำหนดค่าของตัวแปร (2)

สมมติว่าคุณมีตะกร้าผลไม้ซึ่งมีแอปเปิ้ลอยู่ 5 ผล ในฐานะนักวิเคราะห์ข้อมูล คุณต้องการจะเก็บค่าของจำนวนแอปเปิ้ลไว้ในตัวแปรชื่อ `my_apples`

*** =instructions
- พิมพ์ `my_apples <- 5` ลงไปใน editor เพื่อกำหนดให้ตัวแปร `my_apples` มีค่าเป็น 5
- พิมพ์ `my_apples` ลงไปใต้คอมเมนต์ที่สอง เพื่อทำการแสดงค่าของตัวแปร `my_apples`.
- กดปุ่ม 'Submit Answer' คุณจะเห็นว่ามีการแสดงเลข 5 อยู่บน Console นี่เป็นตัวอย่างเพื่อแสดงให้เห็นว่า `my_apples` มีค่าเท่ากับ 5 จริงๆ

*** =hint
จำไว้ว่าถ้าคุณต้องการกำหนดค่าให้กับตัวแปรบนโปรแกรม R คุณสามารถใช้คำสั่งกำหนดค่า `<-` หรือ `=` ก็ได้
แต่ `<-` จะเป็นที่นิยมมากกว่าในหมู่ผู้ใช้ R ด้วยกัน

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Assign the value 5 to the variable my_apples


# Print out the value of the variable my_apples

```

*** =solution
```{r}
# Assign the value 5 to the variable my_apples
my_apples <- 5

# Print out the value of the variable my_apples
my_apples
```

*** =sct
```{r}
undefined_msg <- paste("&#3605;&#3619;&#3623;&#3592;&#3626;&#3629;&#3610;&#3651;&#3627;&#3657;&#3649;&#3609;&#3656;&#3651;&#3592;&#3623;&#3656;&#3634;&#3588;&#3640;&#3603;&#3652;&#3604;&#3657;&#3607;&#3635;&#3585;&#3634;&#3619;&#3585;&#3635;&#3627;&#3609;&#3604;&#3588;&#3656;&#3634;&#3651;&#3627;&#3657;&#3605;&#3633;&#3623;&#3649;&#3611;&#3619;", "`my_apples`", "&#3649;&#3621;&#3657;&#3623;")
# ตรวจสอบให้แน่ใจว่าคุณได้ทำการกำหนดค่าให้ตัวแปร `my_apples` แล้ว

incorrect_msg <- paste("&#3605;&#3619;&#3623;&#3592;&#3626;&#3629;&#3610;&#3651;&#3627;&#3657;&#3649;&#3609;&#3656;&#3651;&#3592;&#3623;&#3656;&#3634;&#3588;&#3640;&#3603;&#3652;&#3604;&#3657;&#3607;&#3635;&#3585;&#3634;&#3619;&#3585;&#3635;&#3627;&#3609;&#3604;&#3588;&#3656;&#3634;&#3607;&#3637;&#3656;&#3606;&#3641;&#3585;&#3605;&#3657;&#3629;&#3591;&#3651;&#3627;&#3657;&#3585;&#3633;&#3610;&#3605;&#3633;&#3623;&#3649;&#3611;&#3619;", "`my_apples`", "&#3649;&#3621;&#3657;&#3623;")
# ตรวจสอบให้แน่ใจว่าคุณได้ทำการกำหนดค่าที่ถูกต้องให้กับตัวแปร `my_apples` แล้ว

incorrect_msg <- paste("&#3588;&#3640;&#3603;&#3652;&#3604;&#3657;&#3610;&#3629;&#3585;&#32;&#82;&#32;&#3651;&#3627;&#3657;&#3649;&#3626;&#3604;&#3591;&#3588;&#3656;&#3634;&#3586;&#3629;&#3591;&#3605;&#3633;&#3623;&#3649;&#3611;&#3619;", "`my_apples`", "&#3629;&#3629;&#3585;&#3617;&#3634;&#3627;&#3619;&#3639;&#3629;&#3618;&#3633;&#3591;&#63;")
# คุณได้บอก R ให้แสดงค่าของตัวแปร `my_apples` ออกมาหรือยัง?

success_msg <- "&#3648;&#3618;&#3637;&#3656;&#3618;&#3617;&#33;&#32;&#3652;&#3611;&#3649;&#3610;&#3610;&#3613;&#3638;&#3585;&#3627;&#3633;&#3604;&#3605;&#3656;&#3629;&#3652;&#3611;&#3652;&#3604;&#3657;&#3648;&#3621;&#3618;&#33;"
# เยี่ยม! ไปแบบฝึกหัดต่อไปได้เลย!

test_object("my_apples", undefined_msg = undefined_msg, incorrect_msg = incorrect_msg)
test_output_contains("my_apples", incorrect_msg = incorrect_msg)
success_msg(success_msg)
```


--- type:NormalExercise xp:100 skills:1 key:1c1bd25045
## การกำหนดค่าของตัวแปร (3)

ต่อจากแบบฝึกหัดที่แล้ว คุณเพิ่มส้มอีก 6 ผลลงไปในตะกร้าเดิมพร้อมทำการสร้างตัวแปร `my_orange` และกำหนดค่า 6 ให้กับมัน คุณต้องการจะหาว่าในตะกร้าตอนนี้มีผลไม้ทั้งหมดกี่ผล เนื่องจากคุณมีตัวแปรที่ถูกตั้งชื่อไว้พร้อมอยู่แล้ว ดังนั้นคุณสามารถหาคำตอบได้โดยการพิมพ์โค้ดว่า

```
my_apples + my_oranges
```

*** =instructions
- กำหนดให้ `my_oranges` มีค่าเป็น 6.
- นำตัวแปร `my_apples` และ `my_oranges` มาบวกกัน แล้วให้ R แสดงค่าผลลัพธ์.
- เก็บค่าผลบวกระหว่าง `my_apples` และ `my_oranges` ไว้ในตัวแปรใหม่ชื่อว่า `my_fruit`.

*** =hint
`my_fruit` ก็คือผลบวกระหว่าง `my_apples` และ `my_oranges` คุณสามารถใช้เครื่องหมาย `+` เพื่อรวมตัวแปรทั้งสองเข้าด้วยกันและใช้คำสั่ง `<-` เพื่อกำหนดค่าของผลลัพธ์จากการบวกให้กับตัวแปร `my_fruit`.

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Assign a value to the variables my_apples and my_oranges
my_apples <- 5


# Add these two variables together


# Create the variable my_fruit

```

*** =solution
```{r}
# Assign a value to the variables my_apples and my_oranges
my_apples  <- 5
my_oranges <- 6

# Add these two variables together
my_apples + my_oranges

# Create the variable my_fruit
my_fruit <- my_apples + my_oranges
```

*** =sct
```{r}
apple_incorrect <- paste("&#3629;&#3618;&#3656;&#3634;&#3621;&#3610;&#3610;&#3619;&#3619;&#3607;&#3633;&#3604;&#3607;&#3637;&#3656;&#3585;&#3635;&#3627;&#3609;&#3604;&#3651;&#3627;&#3657;", "`my_apples`", "&#3617;&#3637;&#3588;&#3656;&#3634;&#3648;&#3611;&#3655;&#3609;&#32;&#53;")
# อย่าลบบรรทัดที่กำหนดให้ `my_apples` มีค่าเป็น 5

orange_incorrect <- paste("&#3629;&#3618;&#3656;&#3634;&#3621;&#3639;&#3617;&#3585;&#3635;&#3627;&#3609;&#3604;&#3651;&#3627;&#3657;", "`my_oranges`", "&#3617;&#3637;&#3588;&#3656;&#3634;&#3648;&#3611;&#3655;&#3609;&#32;&#54;&#32;&#3604;&#3657;&#3623;&#3618;")
# อย่าลืมกำหนดให้ my_oranges มีค่าเป็น 6 ด้วย

output_incorrect <- paste("&#3629;&#3618;&#3656;&#3634;&#3621;&#3639;&#3617;&#3649;&#3626;&#3604;&#3591;&#3588;&#3656;&#3634;&#3612;&#3621;&#3621;&#3633;&#3614;&#3608;&#3660;&#3585;&#3634;&#3619;&#3610;&#3623;&#3585;&#3619;&#3632;&#3627;&#3623;&#3656;&#3634;&#3591;", "`my_apples`", "&#3585;&#3633;&#3610;", "`my_oranges`", "&#3604;&#3657;&#3623;&#3618;&#33;&#32;&#3605;&#3633;&#3623;&#3629;&#3618;&#3656;&#3634;&#3591;&#3607;&#3634;&#3591;&#3595;&#3657;&#3634;&#3618;&#3617;&#3639;&#3629;&#3652;&#3604;&#3657;&#3649;&#3626;&#3604;&#3591;&#3623;&#3636;&#3608;&#3637;&#3585;&#3634;&#3619;&#3605;&#3629;&#3610;&#3588;&#3635;&#3606;&#3634;&#3617;&#3586;&#3657;&#3629;&#3609;&#3637;&#3657;&#3652;&#3623;&#3657;&#3649;&#3621;&#3657;&#3623;")
# อย่าลืมแสดงค่าผลลัพธ์การบวกระหว่าง `my_apples` กับ `my_oranges` ด้วย! ตัวอย่างทางซ้ายมือได้แสดงวิธีการตอบคำถามข้อนี้ไว้แล้ว

msg <- paste("&#3588;&#3640;&#3603;&#3652;&#3604;&#3657;&#3651;&#3594;&#3657;&#3588;&#3635;&#3626;&#3633;&#3656;&#3591;", "`my_fruit <- my_apples + my_oranges`", "&#3651;&#3609;&#3585;&#3634;&#3619;&#3626;&#3619;&#3657;&#3634;&#3591;&#3605;&#3633;&#3623;&#3649;&#3611;&#3619;", "`my_fruit`", "&#3627;&#3619;&#3639;&#3629;&#3648;&#3611;&#3621;&#3656;&#3634;&#63;")
# คุณได้ใช้คำสั่ง `my_fruit <- my_apples + my_oranges` ในการสร้างตัวแปร `my_fruit` หรือเปล่า?

success_msg <- paste("&#3648;&#3618;&#3637;&#3656;&#3618;&#3617;&#33;&#32;&#3611;&#3619;&#3632;&#3650;&#3618;&#3594;&#3609;&#3660;&#3586;&#3657;&#3629;&#3627;&#3609;&#3638;&#3656;&#3591;&#3586;&#3629;&#3591;&#3585;&#3634;&#3619;&#3588;&#3635;&#3609;&#3623;&#3603;&#3650;&#3604;&#3618;&#3651;&#3594;&#3657;&#3605;&#3633;&#3623;&#3649;&#3611;&#3619;&#3588;&#3639;&#3629;&#32;&#3648;&#3619;&#3634;&#3626;&#3634;&#3617;&#3634;&#3619;&#3606;&#3609;&#3635;&#3605;&#3633;&#3623;&#3649;&#3611;&#3619;&#3648;&#3627;&#3621;&#3656;&#3634;&#3609;&#3633;&#3657;&#3609;&#3585;&#3621;&#3633;&#3610;&#3617;&#3634;&#3651;&#3594;&#3657;&#3651;&#3627;&#3617;&#3656;&#3652;&#3604;&#3657;&#32;&#3606;&#3657;&#3634;&#3588;&#3640;&#3603;&#3648;&#3611;&#3621;&#3637;&#3656;&#3618;&#3609;&#3588;&#3656;&#3634;&#3586;&#3629;&#3591;", "`my_apples`", "&#3648;&#3611;&#3655;&#3609;&#32;&#49;&#50;&#32;&#3649;&#3621;&#3657;&#3623;&#3619;&#3633;&#3609;&#3650;&#3588;&#3657;&#3604;&#3651;&#3627;&#3617;&#3656;&#3629;&#3637;&#3585;&#3588;&#3619;&#3633;&#3657;&#3591;&#32;&#3588;&#3656;&#3634;&#3586;&#3629;&#3591;", "`my_fruit`", "&#3585;&#3655;&#3592;&#3632;&#3606;&#3641;&#3585;&#3648;&#3611;&#3621;&#3637;&#3656;&#3618;&#3609;&#3605;&#3634;&#3617;&#3652;&#3611;&#3604;&#3657;&#3623;&#3618;")
# เยี่ยม! ประโยชน์ข้อหนึ่งของการคำนวณโดยใช้ตัวแปรคือ เราสามารถนำตัวแปรเหล่านั้นกลับมาใช้ใหม่ได้ ถ้าคุณเปลี่ยนค่าของ `my_apples` เป็น 12 แล้วรันโค้ดใหม่อีกครั้ง ค่าของ `my_fruit` ก็จะถูกเปลี่ยนตามไปด้วย

test_object("my_apples", incorrect_msg = apple_incorrect)
test_object("my_oranges", incorrect_msg = orange_incorrect)
test_output_contains("my_apples + my_oranges", incorrect_msg = output_incorrect)
test_object("my_fruit", undefined_msg = msg, incorrect_msg = msg)
success_msg(success_msg)
```

--- type:NormalExercise xp:100 skills:1 key:915fcc7c99
## แอปเปิ้ลกับส้ม

ในแบบฝึกหัดก่อนหน้านี้ คุณเพิ่งใส่ส้มเพิ่มลงไปในตะกร้าแอปเปิ้ลจนผลไม้ทั้งสองชนิดปนกันไปหมด! เนื่องจากตัวแปร `my_apples` และ `my_oranges` ได้เก็บค่าของตัวเลขเอาไว้ เราจึงสามารถใช้เครื่องหมาย `+` ในการหาผลบวกระหว่างตัวแปรทั้งสองได้ (เครื่องหมาย `+` สามารถใช้ได้กับตัวแปรที่มีค่าเป็นตัวเลขเท่านั้น) คราวนี้ ถ้าคุณลองกำหนดค่าของตัวแปร `my_oranges` ให้เป็นตัวอักษรหรือข้อความ(ตามที่เห็นใน editor) แล้วนำมาบวกกับ `my_apples` นั่นแสดงว่าคุณกำลังพยายามจับตัวเลขมาบวกกับตัวอักษร แล้วเก็บค่าไว้ในตัวแปร `my_fruit` ซึ่งเป็นไปไม่ได้

*** =instructions
- กดปุ่ม 'Submit Answer' แล้วลองทำความเข้าใจ error message ว่าทำไมเราจึงไม่สามารถทำแบบนี้ได้
- ปรับเปลี่ยนโค้ดเพื่อบอกให้ R เข้าใจว่ามีส้มอยู่ 6 ผล ดังนั้นในตะกร้าผลไม้จึงมีผลไม้ทั้งหมด 11 ผล

*** =hint
คุณต้องกำหนดให้ตัวแปร `my_oranges` มีค่าเป็น `6` แทนที่จะเป็น `"six"` จะสังเกตได้ว่าการใส่เครื่องหมาย `"` คือการบอกว่า `"six"` เป็นข้อความหรือตัวอักษร

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Assign a value to the variable my_apples
my_apples <- 5 

# Fix the assignment of my_oranges
my_oranges <- "six" 

# Create the variable my_fruit and print it out
my_fruit <- my_apples + my_oranges 
my_fruit
```

*** =solution
```{r}
# Assign a value to the variable my_apples
my_apples <- 5  

# Fix the assignment of my_oranges
my_oranges <- 6

# Create the variable my_fruit and print it out
my_fruit <- my_apples + my_oranges 
my_fruit
```

*** =sct
```{r}
error_msg <- paste("&#3588;&#3640;&#3603;&#3605;&#3657;&#3629;&#3591;&#3585;&#3635;&#3627;&#3609;&#3604;&#3651;&#3627;&#3657;&#3605;&#3633;&#3623;&#3649;&#3611;&#3619;", "`my_oranges`", "&#3617;&#3637;&#3588;&#3656;&#3634;&#3648;&#3611;&#3655;&#3609;&#3605;&#3633;&#3623;&#3648;&#3621;&#3586;&#32;&#3652;&#3617;&#3656;&#3651;&#3594;&#3656;&#3605;&#3633;&#3623;&#3629;&#3633;&#3585;&#3625;&#3619;")
# คุณต้องกำหนดให้ตัวแปร `my_oranges` มีค่าเป็นตัวเลข ไม่ใช่ตัวอักษร

apple_incorrect <- paste("`my_apples`", "&#3605;&#3657;&#3629;&#3591;&#3617;&#3637;&#3588;&#3656;&#3634;&#3648;&#3611;&#3655;&#3609;", "5")
# `my_apples` ต้องมีค่าเป็น 5

oranges_incorrect <- paste("`my_oranges`", "&#3605;&#3657;&#3629;&#3591;&#3617;&#3637;&#3588;&#3656;&#3634;&#3648;&#3611;&#3655;&#3609;", "6")
# `my_oranges` ต้องมีค่าเป็น 6

fruit_incorrect <- paste("`my_fruit`", "&#3605;&#3657;&#3629;&#3591;&#3617;&#3637;&#3588;&#3656;&#3634;&#3648;&#3611;&#3655;&#3609;", "11", "&#40;&#3612;&#3621;&#3619;&#3623;&#3617;&#3619;&#3632;&#3627;&#3623;&#3656;&#3634;&#3591;", "`my_apples`", "&#3649;&#3621;&#3632;", "`my_oranges`)")
# `my_fruit` ต้องมีค่าเป็น 11 (ผลรวมระหว่าง `my_apples` และ `my_oranges`)

output_incorrect <- paste("&#3629;&#3618;&#3656;&#3634;&#3621;&#3610;&#3610;&#3619;&#3619;&#3607;&#3633;&#3604;&#3607;&#3637;&#3656;&#3626;&#3633;&#3656;&#3591;&#3651;&#3627;&#3657;&#3649;&#3626;&#3604;&#3591;&#3588;&#3656;&#3634;", "`my_fruit`")
# อย่าลบบรรทัดที่สั่งให้แสดงค่า `my_fruit`

success_msg <- "&#3588;&#3640;&#3603;&#3607;&#3635;&#3652;&#3604;&#3657;&#3604;&#3637;&#3607;&#3637;&#3648;&#3604;&#3637;&#3618;&#3623;&#33;&#32;&#3652;&#3611;&#3604;&#3641;&#3649;&#3610;&#3610;&#3613;&#3638;&#3585;&#3627;&#3633;&#3604;&#3605;&#3656;&#3629;&#3652;&#3611;&#3585;&#3633;&#3609;&#3648;&#3621;&#3618;"
# คุณทำได้ดีทีเดียว! ไปดูแบบฝึกหัดต่อไปกันเลย

test_error(incorrect_msg = error_msg)
test_object("my_apples", incorrect_msg = apple_incorrect)
test_object("my_oranges", incorrect_msg = oranges_incorrect)
test_object("my_fruit", incorrect_msg = fruit_incorrect)
test_output_contains("my_fruit", incorrect_msg = output_incorrect)
success_msg(success_msg)
```

--- type:NormalExercise xp:100 skills:1 key:0f23107394
## ประเภทของข้อมูลพื้นฐานใน R

ในการใช้ R เราจะต้องทำงานกับข้อมูลหลากหลายประเภท เช่น:

- ตัวเลขทศนิยมต่างๆ เช่น `4.5` เรียกว่า **numerics**.
- ตัวเลขจำนวนเต็ม เช่น `4` เรียกว่า **integers** (ซึ่งนับเป็นประเภทหนึ่งของ **numerics** ด้วย).
- ข้อมูลแบบ Boolean หรือตรรกะ (`TRUE` หรือ `FALSE`) เรียกว่า **logical**.
- ข้อความและตัวอักษร (string) เรียกว่า **characters**.

ลองดูในหน้าต่าง editor จะเห็นว่าการใช้เครื่องหมาย `"` คือการบอกว่า "some text" เป็นข้อมูลตัวอักษร

*** =instructions
ลองเปลี่ยนค่าของตัวแปรใน editor ให้ถูกต้องตามรายละเอียดดังต่อไปนี้:

- `my_numeric` ให้มีค่าเป็น `42`.
- `my_character` ให้มีค่าเป็น `"universe"` (เครื่องหมาย `"` เป็นการบอกว่า `"universe"` เป็นข้อมูลตัวอักษร)
- `my_logical` ให้มีค่าเป็น `FALSE`.

จำไว้ว่า R เป็นโปรแกรมที่ case sensitive! (คือการที่อักษรภาษาอังกฤษตัวพิมพ์เล็กและตัวพิมพ์ใหญ่มีความแตกต่างกัน เช่น R จะถือว่า `a` ไม่เท่ากับ `A` เป็นต้น
ดังนั้นในกรณีนี้ `"universe"` จะไม่เท่ากับ `"Universe"`)

*** =hint 
สิ่งที่คุณต้องทำคือกำหนดให้กับตัวแปรมีค่าตรงกับที่กำหนดไว้ใน instructions เช่น `my_numeric <- 42` คือการกำหนดให้ `my_numeric` มีค่าเป็น 42

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Change my_numeric to be 42
my_numeric <- 42.5

# Change my_character to be "universe"
my_character <- "some text"

# Change my_logical to be FALSE
my_logical <- TRUE
```

*** =solution
```{r}
# Change my_numeric to be 42
my_numeric <- 42

# Change my_character to be "universe"
my_character <- "universe"

# Change my_logical to be FALSE
my_logical <- FALSE
```

*** =sct
```{r}
numeric_incorrect <- paste("&#3605;&#3619;&#3623;&#3592;&#3626;&#3629;&#3610;&#3651;&#3627;&#3657;&#3649;&#3609;&#3656;&#3651;&#3592;&#3623;&#3656;&#3634;&#3588;&#3640;&#3603;&#3652;&#3604;&#3657;&#3648;&#3611;&#3621;&#3637;&#3656;&#3618;&#3609;&#3588;&#3656;&#3634;", "`my_numeric`", "&#3648;&#3611;&#3655;&#3609;&#32;&#52;&#50;&#32;&#3649;&#3621;&#3657;&#3623;")
# ตรวจสอบให้แน่ใจว่าคุณได้เปลี่ยนค่า `my_numeric` เป็น 42 แล้ว

char_incorrect <- paste("&#3605;&#3619;&#3623;&#3592;&#3626;&#3629;&#3610;&#3651;&#3627;&#3657;&#3649;&#3609;&#3656;&#3651;&#3592;&#3623;&#3656;&#3634;&#3588;&#3640;&#3603;&#3652;&#3604;&#3657;&#3648;&#3611;&#3621;&#3637;&#3656;&#3618;&#3609;&#3588;&#3656;&#3634;", "`my_character`", "&#3648;&#3611;&#3655;&#3609;", "`\"universe\"`", "&#3649;&#3621;&#3657;&#3623;&#32;&#3629;&#3618;&#3656;&#3634;&#3621;&#3639;&#3617;&#3648;&#3588;&#3619;&#3639;&#3656;&#3629;&#3591;&#3627;&#3617;&#3634;&#3618;", "`\"`", "&#3604;&#3657;&#3623;&#3618;&#3621;&#3656;&#3632;") 
# แน่ใจนะว่าคุณได้เปลี่ยนค่าของ `my_character` เป็น `\"universe\"` แล้ว? อย่าลืมเครื่องหมาย `\"` ด้วยล่ะ

logical_incorrect <- paste("&#3605;&#3619;&#3623;&#3592;&#3626;&#3629;&#3610;&#3651;&#3627;&#3657;&#3649;&#3609;&#3656;&#3651;&#3592;&#3623;&#3656;&#3634;&#3588;&#3640;&#3603;&#3652;&#3604;&#3657;&#3648;&#3611;&#3621;&#3637;&#3656;&#3618;&#3609;&#3588;&#3656;&#3634;", "`my_logical`", "&#3648;&#3611;&#3655;&#3609;", "`FALSE` (`FALSE`", "&#3651;&#3609;&#3607;&#3637;&#3656;&#3609;&#3637;&#3657;&#3588;&#3623;&#3619;&#3648;&#3611;&#3655;&#3609;&#3605;&#3633;&#3623;&#3614;&#3636;&#3617;&#3614;&#3660;&#3651;&#3627;&#3597;&#3656;&#3607;&#3633;&#3657;&#3591;&#3627;&#3617;&#3604;&#3604;&#3657;&#3623;&#3618;&#41;")
# แน่ใจนะว่าคุณได้เปลี่ยนค่าของ `my_logical` เป็น `FALSE` (`FALSE` ในที่นี้ควรเป็นตัวพิมพ์ใหญ่ทั้งหมดด้วย)

success_msg <- "&#3607;&#3635;&#3652;&#3604;&#3657;&#3604;&#3637;&#3617;&#3634;&#3585;&#33;&#32;&#3652;&#3611;&#3607;&#3637;&#3656;&#3649;&#3610;&#3610;&#3613;&#3638;&#3585;&#3627;&#3633;&#3604;&#3605;&#3656;&#3629;&#3652;&#3611;&#3585;&#3633;&#3609;&#3648;&#3621;&#3618;"
# ทำได้ดีมาก! ไปที่แบบฝึกหัดต่อไปกันเลย

test_object("my_numeric", incorrect_msg = numeric_incorrect)
test_object("my_character", incorrect_msg = char_incorrect)
test_object("my_logical", incorrect_msg = logical_incorrect)
success_msg(success_msg)
```

--- type:NormalExercise xp:100 skills:1 key:99b549229d
## การแสดงประเภทของข้อมูล

ในแบบฝึกหัดก่อนๆ เมื่อคุณลองใส่ `5 + "six"` ลงไปใน editor คุณจะเห็น error message ว่าเราไม่สามารถทำแบบนั้นกับข้อมูลคนละประเภทได้
แต่คุณสามารถแก้ไขปัญหานี้ได้โดยการตรวจสอบประเภทของข้อมูลก่อนนำไปใช้ โดยการเรียกใช้ function `class()` ตามตัวอย่างทางด้านขวามือ

*** =instructions
เพิ่มเติมโค้ดลงใน editor เพื่อให้ R แสดงประเภทของข้อมูลที่อยู่ในตัวแปร `my_character` และ `my_logical`. 

*** =hint
คุณสามารถดูตัวอย่างได้จากโค้ดที่ทำการแสดงประเภทของข้อมูลที่อยู่ในตัวแปร `my_numeric` แล้วนำไปปรับใช้กับ `my_character` และ `my_logical` ให้ถูกต้อง

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Declare variables of different types
my_numeric <- 42
my_character <- "universe"
my_logical <- FALSE 

# Check class of my_numeric
class(my_numeric)

# Check class of my_character


# Check class of my_logical

```

*** =solution
```{r}
# Declare variables of different types:
my_numeric <- 42
my_character <- "universe"
my_logical <- FALSE

# Check class of my_numeric
class(my_numeric)

# Check class of my_character
class(my_character)

# Check class of my_logical
class(my_logical)
```

*** =sct
```{r}
msg <- "&#3629;&#3618;&#3656;&#3634;&#3648;&#3611;&#3621;&#3637;&#3656;&#3618;&#3609;&#3649;&#3611;&#3621;&#3591;&#3585;&#3634;&#3619;&#3585;&#3635;&#3627;&#3609;&#3604;&#3588;&#3656;&#3634;&#3605;&#3633;&#3623;&#3649;&#3611;&#3619;&#33;"
# อย่าเปลี่ยนแปลงการกำหนดค่าตัวแปร!

success_msg <- "&#3618;&#3636;&#3609;&#3604;&#3637;&#3604;&#3657;&#3623;&#3618;&#33;&#32;&#3588;&#3640;&#3603;&#3652;&#3604;&#3657;&#3607;&#3635;&#3649;&#3610;&#3610;&#3613;&#3638;&#3585;&#3627;&#3633;&#3604;&#3586;&#3629;&#3591;&#3610;&#3607;&#3609;&#3637;&#3657;&#3588;&#3619;&#3610;&#3607;&#3633;&#3657;&#3591;&#3627;&#3617;&#3604;&#3649;&#3621;&#3657;&#3623;&#33;&#32;&#3610;&#3607;&#3605;&#3656;&#3629;&#3652;&#3611;&#3648;&#3619;&#3634;&#3592;&#3632;&#3652;&#3604;&#3657;&#3648;&#3619;&#3637;&#3618;&#3609;&#3619;&#3641;&#3657;&#3617;&#3634;&#3585;&#3586;&#3638;&#3657;&#3609;&#3648;&#3585;&#3637;&#3656;&#3618;&#3623;&#3585;&#3633;&#3610;&#3626;&#3636;&#3656;&#3591;&#3607;&#3637;&#3656;&#3648;&#3619;&#3637;&#3618;&#3585;&#3623;&#3656;&#3634;&#32;&#118;&#101;&#99;&#116;&#111;&#114;&#115;&#33;"
# ยินดีด้วย! คุณได้ทำแบบฝึกหัดของบทนี้ครบทั้งหมดแล้ว! บทต่อไปเราจะได้เรียนรู้มากขึ้นเกี่ยวกับสิ่งที่เรียกว่า vectors!

lapply(c("my_numeric", "my_character", "my_logical"), test_object, undefined_msg = msg, incorrect_msg = msg)

numeric_incorrect <- paste("&#3629;&#3618;&#3656;&#3634;&#3621;&#3610;&#3650;&#3588;&#3657;&#3604;&#3607;&#3637;&#3656;&#3626;&#3633;&#3656;&#3591;&#3649;&#3626;&#3604;&#3591;&#3588;&#3656;&#3634;&#3611;&#3619;&#3632;&#3648;&#3616;&#3607;&#3586;&#3657;&#3629;&#3617;&#3641;&#3621;&#3607;&#3637;&#3656;&#3629;&#3618;&#3641;&#3656;&#3651;&#3609;&#3605;&#3633;&#3623;&#3649;&#3611;&#3619;", "`my_numeric`")
# อย่าลบโค้ดที่สั่งแสดงค่าประเภทข้อมูลที่อยู่ในตัวแปร `my_numeric`

patt <- paste("&#3605;&#3619;&#3623;&#3592;&#3626;&#3629;&#3610;&#3651;&#3627;&#3657;&#3649;&#3609;&#3656;&#3651;&#3592;&#3623;&#3656;&#3634;&#3588;&#3640;&#3603;&#3652;&#3604;&#3657;&#3614;&#3636;&#3617;&#3614;&#3660;", "`class(%1$s)`", "&#3621;&#3591;&#3652;&#3611;&#3648;&#3614;&#3639;&#3656;&#3629;&#3626;&#3633;&#3656;&#3591;&#3649;&#3626;&#3604;&#3591;&#3611;&#3619;&#3632;&#3648;&#3616;&#3607;&#3586;&#3629;&#3591;&#3586;&#3657;&#3629;&#3617;&#3641;&#3621;&#3607;&#3637;&#3656;&#3629;&#3618;&#3641;&#3656;&#3651;&#3609;&#3605;&#3633;&#3623;&#3649;&#3611;&#3619;", "`%1$s`", "&#3648;&#3619;&#3637;&#3618;&#3610;&#3619;&#3657;&#3629;&#3618;&#3649;&#3621;&#3657;&#3623;")
# ตรวจสอบให้แน่ใจว่าคุณได้พิมพ์ `class(%1$s)` ลงไปเพื่อสั่งแสดงประเภทของข้อมูลที่อยู่ใน `%1$s` เรียบร้อยแล้ว

test_output_contains("class(my_numeric)", incorrect_msg = numeric_incorrect)
test_output_contains("class(my_character)", incorrect_msg = sprintf(patt, "my_character"))
test_output_contains("class(my_logical)", incorrect_msg = sprintf(patt, "my_logical"))
success_msg(success_msg)
```

